# Activity: Using Conditionals

## Introduction

In this activity, you will use conditionals to solve the following questions.
This activity contains:
- `if` statements
- `if-else` statements
- `if-elif-else` statements
- Nested `if` statements
- `if` statements with logical operators


#### Question 1

Using the below modulo operator(%, which returns the remainder after division), write an `if-else` statement that assigns the value `True` to the variable `is_even` if `number` is even or `False` if `number` is odd.

```python
number = 24
is_even = False

#Modulo Operator
number % 2 == 0
```

Output:
```
True
```

<!-- grade_id: cell-182243ac3ee299ba, points: 0, grade: false, solution: true -->
```python
is_even = number % 2 == 0
```

<!-- grade_id: cell-a4bae717c9a2ceff, points: 1, grade: true, solution: false -->
```python
# Question 1 Grading Checks

assert type(is_even) == bool
```

#### Question 2

Write an `if-elif-else` statement that assigns the string `"positive"` to the variable `sign` if the `integer_number` is positive, `"negative"` if `integer_number` is negative, or `"zero"` if `integer_number` is 0.

```python
integer_number = 0
sign = ''
```

<!-- grade_id: cell-45bdb87006d0c115, points: 0, grade: false, solution: true -->
```python
sign = "positive" if integer_number > 0 else "negative" if integer_number < 0 else "zero"
```

<!-- grade_id: cell-23106579aad523d4, points: 1, grade: true, solution: false -->
```python
# Question 2 Grading Checks

assert type(sign) == str
assert len(sign) > 0

```

#### Question 3

You are given a `dict` called `person` with a person's name and an age.

Write a nested `if` statement that first checks if the person is old enough to vote, and assigns the value `True` to the variable `can_vote` if they are older than 17.

Then, check if they are 21 or older and assigns the value `True` to the `can_rent_car` variable.

Finally, print the `can_vote` and `can_rent_car` variables.


```python
person = {
    "name": "James Dean",
    "age": 19,
}
can_vote = False
can_rent_car = False
```

<!-- grade_id: cell-f5ffa30c811c147c, points: 0, grade: false, solution: true -->
```python
can_vote = person["age"] > 17
can_rent_a_car = person["age"] >= 21
```

<!-- grade_id: cell-ced09616fccaa86f, points: 1, grade: true, solution: false -->
```python
# Question 3 Grading Checks

assert type(can_vote) == bool
assert type(can_rent_car) == bool

```

```python

```